---
title: Module 1
description: >-
  This is a use case for the Howard R Class blah blah


---
## Ex 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 26c9007dc9
```

Welcome to Exploratory Data Analysis in R for Howard University Students! If you're here, you should have already completed at least 3 of the Intro to R Programming Datacamp excersizes, as well as the Introduction to the Tidyverse 1: Data Wrangling. If you haven't yet, please follow this link to try them out: https://www.datacamp.com/courses/tech:r

`@instructions`
- Today we're going to be working with _ dataset. 
- Load the dataset using `read_csv()`
- Instruction 3
- Instruction 4

`@hint`
- Here is the hint for this setup problem. It should get students 50% of the way to the correct answer.

`@pre_exercise_code`
```{r}
x <- 1
y <- 2
```
`@sample_code`
```{r}
# Your
x <- 1
y <- 2
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```
`@solution`
```{r}
z <- x + y
# to help students see the differences from solution
# to given.
```
`@sct`
```{r}
test_error()
test_object("z",
            undefined_msg = "Make sure to define `z`!",
            incorrect_msg = "Have you correctly assigned 5 to `z`!")
success_msg("Awesome! It's considered good style to write spaces either side of the assignment arrow.")
```





---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: 2a0bc8f768
```



`@instructions`


`@hint`









