---
title: Module 1
description: >-
  This is a use case for the Howard R Class blah blah


---
## Ex 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 26c9007dc9
```

Welcome to Exploratory Data Analysis in R for Howard University Students! In this chapter we're going to cover some basic concepts about what R is, and get familiar with some basic programming concepts.

`@instructions`
- Type 1 + 1 in the console (that square to the right), and press enter!
- Instruction 2
- Instruction 3
- Instruction 4

`@hint`
- Here is the hint for this setup problem. It should get students 50% of the way to the correct answer.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```
`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```
`@solution`
```{r}
1 + 1
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```
`@sct`
```{r}
# Update this to something more informative.
success_msg("Well done!! Then reinforce a learning objective from the exercise.")
```



